Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel $\rightarrow$ Restart) and then **run all cells** (in the menubar, select Cell $\rightarrow$ Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Vishrut Vanga"
COLLABORATORS = ""
```

---

# **CS 210**: Assignment 1 [100 points]


For this assignment, you will write a Python module called **A1sub.py** in which you will define the seven functions below. 

We are providing you with a stub file **A1sub.py** that contains the function definitions. **Do not change the definitions.** Functions are defined with hints (noted ->) which specify what types the function is expected to output. You can define your own helper functions, as long as the main function for each question follows the given template.

You should submit the Assignment via CodeBench. Click submit next to the assignment. This will submit a snapshot of your whole directory. We will only grade the file **A1sub.py**

You can submit as many times as you want before the deadline, but each submission creates a new snapshot and **erases the previous submission**.



```python
## pprint.pprint() allows to print Dictionaries in a more user-readable format
import pprint
```

### Question 1 [10 points]
Write a program that reads words from a **filename**, which is given as
a string argument. (We will test your function with different files). It should **return** the words from the file in a list, sorted alphabetically in reverse according to the **last character** (and then the second to last, etc).

For instance, if the file has

```
bell tea  Zebra
apple
yellow
```

Then the output should be

`['tea', 'Zebra', 'apple', 'bell', 'yellow']`

<!-- grade_id: cell-91029b5f4c39aa00, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
def reverse_sorted_words(filename: str) -> list[str]:
    text1 = open(filename)
    contents = text1.readlines()
    newlist = []
    for i in range(0, len(contents)):
        contents[i] = contents[i].strip('\n')
        contents[i] = contents[i].split(' ')
    for i in range(0, len(contents)):
        newlist.extend(contents[i])
    newlist = sorted(newlist, key= lambda x: x[::-1])
    return newlist

reverse_sorted_words("words.txt")
```

Output:
```
['tangible',
 'advise',
 'continue',
 'five',
 'fang',
 'whispering',
 'frog',
 'launch',
 'Ask',
 'Thrill',
 'utopian',
 'pan',
 'visitor',
 'houses',
 'point',
 'airport',
 'ray',
 'hurry',
 'activity',
 'buzz']
```

Example output for 

`reverse_sorted_words("words.txt")`

```
['tangible', 'advise', 'continue', 'five', 'fang', 'whispering', 'frog', 'launch', 'Ask', 'Thrill', 'utopian', 'pan', 'visitor', 'houses', 'point', 'airport', 'ray', 'hurry', 'activity', 'buzz']
```

<!-- grade_id: cell-d994a1dde895f631, points: 0, grade: true, solution: false -->
```python
# Use this code block to test your program, please make your own tests too.
from A1sub import reverse_sorted_words
answer_key = ['tangible', 'advise', 'continue', 'five', 'fang', 'whispering', 'frog', 'launch', 'Ask', 'Thrill', 'utopian', 'pan', 'visitor', 'houses', 'point', 'airport', 'ray', 'hurry', 'activity', 'buzz']
rev_words = reverse_sorted_words("words.txt")
try:
    assert len(answer_key) == len(rev_words), "Incorrect length"
    for ans, res in zip(answer_key, rev_words):
        # check if each word matches the answer key
        assert ans == res, str(res) + " not correct"
except AssertionError as e:
        print(f"Test Failed: {e}")
else:
    print("Test Passed")
    
```

Output:
```
Test Passed
```

### Question 2 [15 points]

Your professor asked your TA to collect student information and store it in a file called **studentInfo.txt**, but this TA spilled Pepsi on the keyboard while entering the data. As a result, there are extra or missing spaces and missing commas throughout the file. For example, the file has

```
wd50,WillDavidson, 090,3
hpj4332, Helen              Jordan1983.88
xd222XavierDavies1983.65
ad12  AdDavies1983.65
```
Suppose the file contains netID, name, major, and GPA in that order. Can you help this TA make the file more readable? 

To accomplish this task, you may assume the following about each field: 
`netID`: begins with 2-3 lowercase letters followed by 1-4 digits (e.g., wd1, wxo001, tsy1234)
`name`: contains both first name and last name with their first letter capitalized (e.g., "John   Doe", "MarySmith"). There may be zero or more spaces between the first name and last name. 
`major`: 3-digit code (e.g., 198)
`GPA`: floating point number between 0.0 and 4.0 (inclusive). There may or may not be a decimal point (e.g., 3.0 or 3). 

Write a function called **reformat_student_info** to convert a file into a standard format, which has exactly one comma followed by one space to separate each field. The function should accept a **filename** string parameter. Write the cleaned data to a new file called **whose name is a parameter of your function: outputfilename** and format the file as follows:

`netID, name, major, GPA`

There should be exactly one comma followed by one space to separate each field. The name field should include exactly one space between first name and last name, and the GPA should have a decimal even if the GPA from the input file does not (e.g., 3.0 instead of 3). 

The reformatted file for the sample above should look like this: 

```
wd50, Will Davidson, 090, 3.0
hpj4332, Helen Jordan, 198, 3.88
xd222, Xavier Davies, 198, 3.65
ad12, Ad Davies, 198, 3.65
```
 

<!-- grade_id: cell-febddcc2fe833c4e, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
def reformat_student_info(filename: str, outputfilename: str) -> None:
    ourfile = open(outputfilename, "w")
    for i, line in enumerate(open(filename)):
        flds = line.split()
        for i in range(len(flds)):
            flds[i] = flds[i].strip(",")
        flds = "".join(flds)
        for i in range(len(flds)):
            flds = flds.replace(",", "")
        for i in range(len(flds)):
            if flds[i] >= 'A' and flds[i] <= 'Z':
                netidfield = flds[0:i]
                #print(netidfield, end=" ")
                break
        flds = flds.strip(netidfield)
        if "." in flds:
            for i in range(len(flds)):
                if flds[i] == ".":
                    gpa = flds[i-1:]
            #print(gpa, end=" ")
        else:
            gpa = flds[-1] + '.0'
            #print(gpa, end=" ")
        if "." in flds:
            flds = flds.strip(gpa)
        else:
            flds = flds[:len(flds)-1]
        for i in range(len(flds)-1, 0, -1):
            if flds[i] >= 'a' and flds[i] <= 'z':
                major = flds[i+1:]
                #print(major, end=" ")
                flds = flds.strip(major)
                break
        for i in range(1, len(flds)):
            namestr = ""
            if(flds[i] >= 'A' and flds[i] <= 'Z'):
                firsthalf = flds[0:i]
                secondhalf = flds[i:]
                namestr = firsthalf + " " + secondhalf
                break
        ourfile.write(netidfield + ", " + namestr + ", " + major + ", " + gpa + "\n")
    ourfile.close()
                
        
        
        
        
        
        
        
reformat_student_info('studentInfo.txt','studentInfoReformatted.txt')        
```

```python
# Your cleaned file should look like this
from A1sub import reformat_student_info

reformat_student_info('studentInfo.txt','studentInfoReformatted.txt')
f = open('studentInfoReformatted.txt')
for i, line in enumerate(f):
    print(line.strip())

f.close()
```

Output:
```
ed18, Edward Duran, 077, 3.5
abc1234, Mike Peterson, 12, 3.7
dw0, Danika Wilson, 08, 3.9
wd50, Will Davidson, 090, 3.0
hpj4332, Helen Jordan, 19, 3.88
xd222, Xavier Davies, 198, 3.65
ah1, Allisson Henson, 147, 2.9
ne099, Nina Estrada, 095, 2.0
sh5555, Skylar Huff, 112, 3.3
```

```
ed18, Edward Duran, 077, 3.5
abc1234, Mike Peterson, 123, 3.7
dw0, Danika Wilson, 083, 3.9
wd50, Will Davidson, 090, 3.0
hpj4332, Helen Jordan, 198, 3.88
xd222, Xavier Davies, 198, 3.65
ah1, Allisson Henson, 147, 2.9
ne099, Nina Estrada, 095, 2.0
sh5555, Skylar Huff, 112, 3.3
```

### Question 3  [15 points]
The top_spotify_songs_USA_earlyOct24.csv contains streaming info about the top 50 songs streamed daily on Spotify from Oct 1 to Oct 10. (source:  https://www.kaggle.com/datasets/asaniczka/top-spotify-songs-in-73-countries-daily-updated)

You should look at the data to understand what is represented there. Make sure to understand the format of each column.

Write a function **getDailySongStreamingData** that takes a **filename** as a string parameter (we will test your code with different files) and return a Dictionary that for each song name has a list of the popularity score for the days in which the song appear.  The test file has 10 days, but we will test with files that have more days. 

The Dictionary should look like this:
{'song_name1':[93,91,95,98,81,72,94,86,48],'song_name2':[72,81]}



<!-- grade_id: cell-e2053a97d07532a1, points: 0, grade: true, solution: true -->
```python
from typing import Dict, List
import csv
# YOUR CODE HERE
def getDailySongStreamingData(filename: str) -> Dict[str, List[str]]:
    a = {}
    with open(filename) as file:
        reader = csv.reader(file)
        next(reader)
        for num, row in enumerate(reader):
            if row[1] not in a.keys():
                a[row[1]] = [row[8]]
            else:
                a[row[1]].append(row[8])
        newa = list(a.keys())
        newa.sort()
        sd = {i: a[i] for i in newa}
    return sd
    
getDailySongStreamingData('top_spotify_songs_USA_earlyOct24.csv')
    
```

Output:
```
{'28': ['84', '84', '84', '85', '85', '85', '85', '85', '85', '85'],
 '505': ['80', '80'],
 'A Bar Song (Tipsy)': ['86',
  '86',
  '86',
  '86',
  '86',
  '86',
  '86',
  '86',
  '86',
  '87'],
 'ALL RED': ['90', '90', '91', '91'],
 "Ain't No Love In Oklahoma (From Twisters: The Album)": ['85',
  '85',
  '85',
  '85',
  '85',
  '85',
  '85',
  '85',
  '85',
  '85'],
 "Austin (Boots Stop Workin')": ['89', '90', '90', '90'],
 'BIRDS OF A FEATHER': ['98',
  '98',
  '98',
  '98',
  '98',
  '98',
  '98',
  '98',
  '99',
  '99'],
 'Beautiful Things': ['90',
  '90',
  '90',
  '90',
  '90',
  '90',
  '90',
  '90',
  '90',
  '90'],
 'Bed Chem': ['92', '92', '92', '92', '92', '92', '92', '92', '92', '92'],
 'Big Dawgs': ['91'],
 'Casual': ['86', '86', '86', '86', '86', '86', '86', '86'],
 'Cowgirls (feat. ERNEST)': ['84',
  '84',
  '84',
  '84',
  '84',
  '84',
  '85',
  '85',
  '85',
  '85'],
 'Die With A Smile': ['100',
  '100',
  '100',
  '100',
  '100',
  '100',
  '100',
  '100',
  '100',
  
```

```python
# Your output should look like this 
from A1sub import getDailySongStreamingData
song_popularities = getDailySongStreamingData('top_spotify_songs_USA_earlyOct24.csv')
pprint.pprint(song_popularities, width=120)
```

Output:
```
{'28': ['84', '84', '84', '85', '85', '85', '85', '85', '85', '85'],
 '505': ['80', '80'],
 'A Bar Song (Tipsy)': ['86', '86', '86', '86', '86', '86', '86', '86', '86', '87'],
 'ALL RED': ['90', '90', '91', '91'],
 "Ain't No Love In Oklahoma (From Twisters: The Album)": ['85', '85', '85', '85', '85', '85', '85', '85', '85', '85'],
 "Austin (Boots Stop Workin')": ['89', '90', '90', '90'],
 'BIRDS OF A FEATHER': ['98', '98', '98', '98', '98', '98', '98', '98', '99', '99'],
 'Beautiful Things': ['90', '90', '90', '90', '90', '90', '90', '90', '90', '90'],
 'Bed Chem': ['92', '92', '92', '92', '92', '92', '92', '92', '92', '92'],
 'Big Dawgs': ['91'],
 'Casual': ['86', '86', '86', '86', '86', '86', '86', '86'],
 'Cowgirls (feat. ERNEST)': ['84', '84', '84', '84', '84', '84', '85', '85', '85', '85'],
 'Die With A Smile': ['100', '100', '100', '100', '100', '100', '100', '100', '100', '100'],
 'Diet Pepsi': ['87', '87', '87', '87', '86', '86', '86', '86'],
 'Espresso': ['86', '86', '86', '86
```

```
{'28': ['84', '84', '84', '85', '85', '85', '85', '85', '85', '85'],
 '505': ['80', '80'],
 'A Bar Song (Tipsy)': ['86', '86', '86', '86', '86', '86', '86', '86', '86', '87'],
 'ALL RED': ['90', '90', '91', '91'],
 "Ain't No Love In Oklahoma (From Twisters: The Album)": ['85', '85', '85', '85', '85', '85', '85', '85', '85', '85'],
 "Austin (Boots Stop Workin')": ['89', '90', '90', '90'],
 'BIRDS OF A FEATHER': ['98', '98', '98', '98', '98', '98', '98', '98', '99', '99'],
 'Beautiful Things': ['90', '90', '90', '90', '90', '90', '90', '90', '90', '90'],
 'Bed Chem': ['92', '92', '92', '92', '92', '92', '92', '92', '92', '92'],
 'Big Dawgs': ['91'],
 'Casual': ['86', '86', '86', '86', '86', '86', '86', '86'],
 'Cowgirls (feat. ERNEST)': ['84', '84', '84', '84', '84', '84', '85', '85', '85', '85'],
 'Die With A Smile': ['100', '100', '100', '100', '100', '100', '100', '100', '100', '100'],
 'Diet Pepsi': ['87', '87', '87', '87', '86', '86', '86', '86'],
 'Espresso': ['86', '86', '86', '86', '85', '85', '85', '84', '84', '83'],
 'FE!N (feat. Playboi Carti)': ['89', '89', '89', '89', '89', '89', '89', '89', '90', '90'],
 'Good Graces': ['89', '89', '89', '89', '89', '89', '89', '89', '89', '89'],
 'Good Luck, Babe!': ['96', '96', '96', '96', '96', '96', '96', '96', '96', '96'],
 'Guess featuring Billie Eilish': ['90', '90', '73', '91', '91', '91', '91', '91', '91', '91'],
 'HOT TO GO!': ['90', '90', '91', '91', '91', '91', '91', '91', '91', '91'],
 'Heading South': ['84', '84', '84', '84', '84', '84', '84', '84', '84', '85'],
 'I Had Some Help (Feat. Morgan Wallen)': ['85', '85', '85', '85', '85', '85', '85', '85', '86', '86'],
 "I Love You, I'm Sorry": ['89', '89', '89', '89', '89', '89', '89', '89', '89', '89'],
 'I Remember Everything (feat. Kacey Musgraves)': ['87', '87', '87', '87', '87', '87', '87', '87', '88', '88'],
 "It's ok I'm ok": ['86', '87', '87', '86', '86', '86', '86', '86'],
 'Juno': ['88', '88', '88', '88', '88', '88', '88', '88', '88', '88'],
 'LA PATRULLA': ['86', '86', '86', '86', '86', '86', '87', '87', '87', '87'],
 'Last Night': ['81', '81', '81', '82', '82', '82', '82', '82'],
 'Lies Lies Lies': ['84', '84', '84', '84', '85', '85', '85', '85', '85', '85'],
 'Lose Control': ['90', '90', '90', '90', '90', '90', '90', '90', '90', '91'],
 'MILLION DOLLAR BABY': ['91', '91', '91', '91', '91', '91', '91', '91', '91', '91'],
 'Mmhmm': ['80', '80', '80'],
 'My Love Mine All Mine': ['88', '88'],
 'NEL': ['87', '87', '87', '87'],
 'NIGHTS LIKE THIS': ['88', '88', '88', '88', '88', '88', '88', '89', '89', '89'],
 'No One Noticed': ['77', '77', '77', '77', '77', '77'],
 'Not Like Us': ['91', '91', '91', '91', '91', '91', '92', '92', '92', '92'],
 'Or What': ['64', '63', '61'],
 'Pink + White': ['87', '87', '87', '87', '87', '87', '87', '87', '87', '87'],
 'Pink Pony Club': ['88', '88', '88', '88', '88', '88', '88', '88', '88', '88'],
 'Pink Skies': ['87', '87', '87', '87', '87', '87', '87', '87', '87', '87'],
 'Please Please Please': ['86', '86', '86', '85', '85', '85', '85', '84', '84', '84'],
 'Red Wine Supernova': ['86', '87', '87', '87', '87', '87', '87', '87', '87', '87'],
 'Sailor Song': ['84', '83', '82', '81', '80', '79', '76', '73', '69', '67'],
 'See You Again (feat. Kali Uchis)': ['89', '89', '89', '89', '89', '89', '89', '89', '89', '89'],
 'Slow It Down': ['88', '88', '88', '88', '88', '88', '88', '88', '89', '89'],
 'Something in the Orange': ['86', '86', '86', '86', '86', '86', '86', '86', '86', '87'],
 'Stick Season': ['88', '88', '88', '88', '88', '88', '88', '88', '88', '88'],
 'Sun to Me': ['83', '83', '83', '83', '83', '84', '84', '84', '84'],
 'Sweater Weather': ['91', '91', '91', '91', '91', '91', '91', '91', '91', '91'],
 'Taste': ['95', '95', '95', '95', '95', '95', '95', '95', '95', '96'],
 'Thinkin’ Bout Me': ['83', '83'],
 'Timeless (with Playboi Carti)': ['90', '90', '90', '89', '89', '88', '88', '87', '86', '84'],
 'Too Sweet': ['77', '77', '81', '77', '77', '76', '76', '76', '75', '75'],
 'WILDFLOWER': ['91', '91', '91', '91', '91', '91', '91', '91', '91', '91'],
 'Who': ['93', '93', '93', '93', '93', '93', '93', '93', '93', '93'],
 'Wind Up Missin’ You': ['44', '39'],
 'one of wun': ['85', '85', '85', '85', '85', '85', '85', '85', '85', '85'],
 'we fell in love in october': ['86', '86', '86', '85', '85', '85']}
```

### Question 4 [10 points]
Based on the same top_spotify_songs_USA_earlyOct24.csv file

Write a function **getTopSongsInfo** that takes a **filename** as a string parameter (we will test your code with different files) and return a Dictionary that for each song name has the name of the artist(s) in a list. Make sure to strip whitespaces at the beginning and end of the artists' names.

The Dictionary should look like this:
{'song_name1':['artist_name1'],'song_name2':['artist_name2a','artist_name2b']}

```python
#FUNCTION FOOTPRINT
#def getTopSongsInfo(filename):
# returns dict
def getTopSongsInfo(filename: str) -> Dict[str, List[str]]:
    a = {}
    with open(filename) as file:
        reader = csv.reader(file)
        next(reader)
        for num, row in enumerate(reader):
            if row[1] not in a.keys():
                if "," not in row[2]:
                    a[row[1]] = [row[2]]
                else:
                    rownew = row[2].split(", ")
                    a[row[1]] = rownew
        newa = list(a.keys())
        newa.sort()
        sd = {i: a[i] for i in newa}
    return sd


getTopSongsInfo('top_spotify_songs_USA_earlyOct24.csv')
```

Output:
```
{'28': ['Zach Bryan'],
 '505': ['Arctic Monkeys'],
 'A Bar Song (Tipsy)': ['Shaboozey'],
 'ALL RED': ['Playboi Carti'],
 "Ain't No Love In Oklahoma (From Twisters: The Album)": ['Luke Combs'],
 "Austin (Boots Stop Workin')": ['Dasha'],
 'BIRDS OF A FEATHER': ['Billie Eilish'],
 'Beautiful Things': ['Benson Boone'],
 'Bed Chem': ['Sabrina Carpenter'],
 'Big Dawgs': ['Hanumankind', 'Kalmi'],
 'Casual': ['Chappell Roan'],
 'Cowgirls (feat. ERNEST)': ['Morgan Wallen', 'ERNEST'],
 'Die With A Smile': ['Lady Gaga', 'Bruno Mars'],
 'Diet Pepsi': ['Addison Rae'],
 'Espresso': ['Sabrina Carpenter'],
 'FE!N (feat. Playboi Carti)': ['Travis Scott', 'Playboi Carti'],
 'Good Graces': ['Sabrina Carpenter'],
 'Good Luck, Babe!': ['Chappell Roan'],
 'Guess featuring Billie Eilish': ['Charli xcx', 'Billie Eilish'],
 'HOT TO GO!': ['Chappell Roan'],
 'Heading South': ['Zach Bryan'],
 'I Had Some Help (Feat. Morgan Wallen)': ['Post Malone', 'Morgan Wallen'],
 "I Love You, I'm Sorry": ['Gracie Abrams'],
 
```

```python
# Your output should look like this 
from A1sub import getTopSongsInfo
top_songs = getTopSongsInfo('top_spotify_songs_USA_earlyOct24.csv')
pprint.pprint(top_songs, width=100)
```

Output:
```
{'28': ['Zach Bryan'],
 '505': ['Arctic Monkeys'],
 'A Bar Song (Tipsy)': ['Shaboozey'],
 'ALL RED': ['Playboi Carti'],
 "Ain't No Love In Oklahoma (From Twisters: The Album)": ['Luke Combs'],
 "Austin (Boots Stop Workin')": ['Dasha'],
 'BIRDS OF A FEATHER': ['Billie Eilish'],
 'Beautiful Things': ['Benson Boone'],
 'Bed Chem': ['Sabrina Carpenter'],
 'Big Dawgs': ['Hanumankind', 'Kalmi'],
 'Casual': ['Chappell Roan'],
 'Cowgirls (feat. ERNEST)': ['Morgan Wallen', 'ERNEST'],
 'Die With A Smile': ['Lady Gaga', 'Bruno Mars'],
 'Diet Pepsi': ['Addison Rae'],
 'Espresso': ['Sabrina Carpenter'],
 'FE!N (feat. Playboi Carti)': ['Travis Scott', 'Playboi Carti'],
 'Good Graces': ['Sabrina Carpenter'],
 'Good Luck, Babe!': ['Chappell Roan'],
 'Guess featuring Billie Eilish': ['Charli xcx', 'Billie Eilish'],
 'HOT TO GO!': ['Chappell Roan'],
 'Heading South': ['Zach Bryan'],
 'I Had Some Help (Feat. Morgan Wallen)': ['Post Malone', 'Morgan Wallen'],
 "I Love You, I'm Sorry": ['Gracie Abrams'],
 
```

```
{'28': ['Zach Bryan'],
 '505': ['Arctic Monkeys'],
 'A Bar Song (Tipsy)': ['Shaboozey'],
 'ALL RED': ['Playboi Carti'],
 "Ain't No Love In Oklahoma (From Twisters: The Album)": ['Luke Combs'],
 "Austin (Boots Stop Workin')": ['Dasha'],
 'BIRDS OF A FEATHER': ['Billie Eilish'],
 'Beautiful Things': ['Benson Boone'],
 'Bed Chem': ['Sabrina Carpenter'],
 'Big Dawgs': ['Hanumankind', 'Kalmi'],
 'Casual': ['Chappell Roan'],
 'Cowgirls (feat. ERNEST)': ['Morgan Wallen', 'ERNEST'],
 'Die With A Smile': ['Lady Gaga', 'Bruno Mars'],
 'Diet Pepsi': ['Addison Rae'],
 'Espresso': ['Sabrina Carpenter'],
 'FE!N (feat. Playboi Carti)': ['Travis Scott', 'Playboi Carti'],
 'Good Graces': ['Sabrina Carpenter'],
 'Good Luck, Babe!': ['Chappell Roan'],
 'Guess featuring Billie Eilish': ['Charli xcx', 'Billie Eilish'],
 'HOT TO GO!': ['Chappell Roan'],
 'Heading South': ['Zach Bryan'],
 'I Had Some Help (Feat. Morgan Wallen)': ['Post Malone', 'Morgan Wallen'],
 "I Love You, I'm Sorry": ['Gracie Abrams'],
 'I Remember Everything (feat. Kacey Musgraves)': ['Zach Bryan', 'Kacey Musgraves'],
 "It's ok I'm ok": ['Tate McRae'],
 'Juno': ['Sabrina Carpenter'],
 'LA PATRULLA': ['Peso Pluma', 'Neton Vega'],
 'Last Night': ['Morgan Wallen'],
 'Lies Lies Lies': ['Morgan Wallen'],
 'Lose Control': ['Teddy Swims'],
 'MILLION DOLLAR BABY': ['Tommy Richman'],
 'Mmhmm': ['BigXthaPlug'],
 'My Love Mine All Mine': ['Mitski'],
 'NEL': ['Fuerza Regida'],
 'NIGHTS LIKE THIS': ['The Kid LAROI'],
 'No One Noticed': ['The Marías'],
 'Not Like Us': ['Kendrick Lamar'],
 'Or What': ['NLE Choppa', '41', 'Kyle Richh'],
 'Pink + White': ['Frank Ocean'],
 'Pink Pony Club': ['Chappell Roan'],
 'Pink Skies': ['Zach Bryan'],
 'Please Please Please': ['Sabrina Carpenter'],
 'Red Wine Supernova': ['Chappell Roan'],
 'Sailor Song': ['Gigi Perez'],
 'See You Again (feat. Kali Uchis)': ['Tyler', 'The Creator', 'Kali Uchis'],
 'Slow It Down': ['Benson Boone'],
 'Something in the Orange': ['Zach Bryan'],
 'Stick Season': ['Noah Kahan'],
 'Sun to Me': ['Zach Bryan'],
 'Sweater Weather': ['The Neighbourhood'],
 'Taste': ['Sabrina Carpenter'],
 'Thinkin’ Bout Me': ['Morgan Wallen'],
 'Timeless (with Playboi Carti)': ['The Weeknd', 'Playboi Carti'],
 'Too Sweet': ['Hozier'],
 'WILDFLOWER': ['Billie Eilish'],
 'Who': ['Jimin'],
 'Wind Up Missin’ You': ['Tucker Wetmore'],
 'one of wun': ['Gunna'],
 'we fell in love in october': ['girl in red']}
```

### Question 5 [15 points]

Write a function **getTopArtistsInfo** that takes a **dictionary** of {title:[list of artists]} similar to the output of question 4 and returns for each artist ***a dictionary** with list of their titles.

The Dictionary should look like this:
{'artist_name1':['song_name1'],'artist_name2':['song_name2a','song_name2b']}

```python
#FUNCTION FOOTPRINT
#def getTopArtistsInfo(dict):
# returns dict
def getTopArtistsInfo(song_dict: Dict[str, List[str]]) -> Dict[str, List[str]]:
    newdict = {}
    for i in song_dict:
        for j in song_dict[i]:
            if j not in newdict.keys():
                newdict[j] = [i]
            else:
                newdict[j].append(i)
    newa = list(newdict.keys())
    newa.sort()
    sd = {i: newdict[i] for i in newa}
    Dict = sd
    return Dict

top_artists = getTopArtistsInfo(top_songs)
print(top_artists)
```

Output:
```
{'41': ['Or What'], 'Addison Rae': ['Diet Pepsi'], 'Arctic Monkeys': ['505'], 'Benson Boone': ['Beautiful Things', 'Slow It Down'], 'BigXthaPlug': ['Mmhmm'], 'Billie Eilish': ['BIRDS OF A FEATHER', 'Guess featuring Billie Eilish', 'WILDFLOWER'], 'Bruno Mars': ['Die With A Smile'], 'Chappell Roan': ['Casual', 'Good Luck, Babe!', 'HOT TO GO!', 'Pink Pony Club', 'Red Wine Supernova'], 'Charli xcx': ['Guess featuring Billie Eilish'], 'Dasha': ["Austin (Boots Stop Workin')"], 'ERNEST': ['Cowgirls (feat. ERNEST)'], 'Frank Ocean': ['Pink + White'], 'Fuerza Regida': ['NEL'], 'Gigi Perez': ['Sailor Song'], 'Gracie Abrams': ["I Love You, I'm Sorry"], 'Gunna': ['one of wun'], 'Hanumankind': ['Big Dawgs'], 'Hozier': ['Too Sweet'], 'Jimin': ['Who'], 'Kacey Musgraves': ['I Remember Everything (feat. Kacey Musgraves)'], 'Kali Uchis': ['See You Again (feat. Kali Uchis)'], 'Kalmi': ['Big Dawgs'], 'Kendrick Lamar': ['Not Like Us'], 'Kyle Richh': ['Or What'], 'Lady Gaga': ['Die With A Smile'], 'Luke Comb
```

```python
from A1sub import getTopArtistsInfo
top_artists = getTopArtistsInfo(top_songs)
pprint.pprint(top_artists, width=150)
```

Output:
```
None
```

```
{'41': ['Or What'],
 'Addison Rae': ['Diet Pepsi'],
 'Arctic Monkeys': ['505'],
 'Benson Boone': ['Beautiful Things', 'Slow It Down'],
 'BigXthaPlug': ['Mmhmm'],
 'Billie Eilish': ['BIRDS OF A FEATHER', 'WILDFLOWER', 'Guess featuring Billie Eilish'],
 'Bruno Mars': ['Die With A Smile'],
 'Chappell Roan': ['Good Luck, Babe!', 'HOT TO GO!', 'Pink Pony Club', 'Red Wine Supernova', 'Casual'],
 'Charli xcx': ['Guess featuring Billie Eilish'],
 'Dasha': ["Austin (Boots Stop Workin')"],
 'ERNEST': ['Cowgirls (feat. ERNEST)'],
 'Frank Ocean': ['Pink + White'],
 'Fuerza Regida': ['NEL'],
 'Gigi Perez': ['Sailor Song'],
 'Gracie Abrams': ["I Love You, I'm Sorry"],
 'Gunna': ['one of wun'],
 'Hanumankind': ['Big Dawgs'],
 'Hozier': ['Too Sweet'],
 'Jimin': ['Who'],
 'Kacey Musgraves': ['I Remember Everything (feat. Kacey Musgraves)'],
 'Kali Uchis': ['See You Again (feat. Kali Uchis)'],
 'Kalmi': ['Big Dawgs'],
 'Kendrick Lamar': ['Not Like Us'],
 'Kyle Richh': ['Or What'],
 'Lady Gaga': ['Die With A Smile'],
 'Luke Combs': ["Ain't No Love In Oklahoma (From Twisters: The Album)"],
 'Mitski': ['My Love Mine All Mine'],
 'Morgan Wallen': ['I Had Some Help (Feat. Morgan Wallen)', 'Cowgirls (feat. ERNEST)', 'Lies Lies Lies', 'Last Night', 'Thinkin’ Bout Me'],
 'NLE Choppa': ['Or What'],
 'Neton Vega': ['LA PATRULLA'],
 'Noah Kahan': ['Stick Season'],
 'Peso Pluma': ['LA PATRULLA'],
 'Playboi Carti': ['Timeless (with Playboi Carti)', 'FE!N (feat. Playboi Carti)', 'ALL RED'],
 'Post Malone': ['I Had Some Help (Feat. Morgan Wallen)'],
 'Sabrina Carpenter': ['Taste', 'Espresso', 'Bed Chem', 'Please Please Please', 'Good Graces', 'Juno'],
 'Shaboozey': ['A Bar Song (Tipsy)'],
 'Tate McRae': ["It's ok I'm ok"],
 'Teddy Swims': ['Lose Control'],
 'The Creator': ['See You Again (feat. Kali Uchis)'],
 'The Kid LAROI': ['NIGHTS LIKE THIS'],
 'The Marías': ['No One Noticed'],
 'The Neighbourhood': ['Sweater Weather'],
 'The Weeknd': ['Timeless (with Playboi Carti)'],
 'Tommy Richman': ['MILLION DOLLAR BABY'],
 'Travis Scott': ['FE!N (feat. Playboi Carti)'],
 'Tucker Wetmore': ['Wind Up Missin’ You'],
 'Tyler': ['See You Again (feat. Kali Uchis)'],
 'Zach Bryan': ['I Remember Everything (feat. Kacey Musgraves)', 'Pink Skies', 'Something in the Orange', '28', 'Heading South', 'Sun to Me'],
 'girl in red': ['we fell in love in october']}
```

### Question 6 [15 points]

Given a **Dictionary with daily popularity scores** for songs in the top charts as produced by Question 3, and a **Dictionary of songs by artists** as produced by Question 6. Write a **getTopArtistsPop** function that computes a dictionary that for each artist gives the average popularity of their top songs. (You need to first average the popularity of a song over the days it appears in the charts (rounded to two decimal places), and then take the average of the song popularity for each artist.


The Dictionary should look like this:
{'artist_name1':avg_popularity1,'artist_name2':avg_popularity2}



```python
#FUNCTION FOOTPRINT
#def getTopArtistsPop(songpopularity_dict, artists_dict):
# returns dict

def getTopArtistsInfo(song_dict: Dict[str, List[str]]) -> Dict[str, List[str]]:
    newdict = {}
    for i in song_dict:
        for j in song_dict[i]:
            if j not in newdict.keys():
                newdict[j] = [i]
            else:
                newdict[j].append(i)
    newa = list(newdict.keys())
    newa.sort()
    sd = {i: newdict[i] for i in newa}
    Dict = sd
    return Dict

top_artists = getTopArtistsInfo(top_songs)

def getTopArtistsPop(song_pop_dict: Dict[str, List[str]], artist_dict: Dict[str, List[str]]) -> Dict[str, float]:
    newdict = {}
    averagedict = {}
    for i in song_pop_dict:
        sum = 0
        for j in song_pop_dict[i]:
            sum += float(j)
        averagedict[i] = float(sum / len(song_pop_dict[i]))
    for i in artist_dict:
        newdict[i] = []
    for i in newdict:
        for j in artist_dict[i]:
            if newdict[i] == []:
                newdict[i] = [averagedict[j]]
            else:
                newdict[i].append(averagedict[j])
    for i in newdict:
        sum = 0
        for j in newdict[i]:
            sum += j
        newavg = (sum/(len(newdict[i])))
        newdict[i] = newavg
    for i in newdict:
        if len(str(newdict[i])) > 5:
            newdict[i] = f"{newdict[i]:.2f}"
            newdict[i] = float(newdict[i])
    Dict = newdict
    return Dict
    
    
    
    
    
            

getTopArtistsPop(song_popularities,top_artists)
```

Output:
```
{'41': 62.67,
 'Addison Rae': 86.5,
 'Arctic Monkeys': 80.0,
 'Benson Boone': 89.1,
 'BigXthaPlug': 80.0,
 'Billie Eilish': 92.73,
 'Bruno Mars': 100.0,
 'Chappell Roan': 89.54,
 'Charli xcx': 89.0,
 'Dasha': 89.75,
 'ERNEST': 84.4,
 'Frank Ocean': 87.0,
 'Fuerza Regida': 87.0,
 'Gigi Perez': 77.4,
 'Gracie Abrams': 89.0,
 'Gunna': 85.0,
 'Hanumankind': 91.0,
 'Hozier': 76.7,
 'Jimin': 93.0,
 'Kacey Musgraves': 87.2,
 'Kali Uchis': 89.0,
 'Kalmi': 91.0,
 'Kendrick Lamar': 91.4,
 'Kyle Richh': 62.67,
 'Lady Gaga': 100.0,
 'Luke Combs': 85.0,
 'Mitski': 88.0,
 'Morgan Wallen': 83.77,
 'NLE Choppa': 62.67,
 'Neton Vega': 86.4,
 'Noah Kahan': 88.0,
 'Peso Pluma': 86.4,
 'Playboi Carti': 89.27,
 'Post Malone': 85.2,
 'Sabrina Carpenter': 89.02,
 'Shaboozey': 86.1,
 'Tate McRae': 86.25,
 'Teddy Swims': 90.1,
 'The Creator': 89.0,
 'The Kid LAROI': 88.3,
 'The Marías': 77.0,
 'The Neighbourhood': 91.0,
 'The Weeknd': 88.1,
 'Tommy Richman': 91.0,
 'Travis Scott': 89.2,
 'Tucker Wetmore': 41.5
```

```python
from A1sub import getTopArtistsPop
artists_pop = getTopArtistsPop(song_popularities,top_artists)
pprint.pprint(artists_pop)
```

Output:
```
None
```

```
{'41': 62.67,
 'Addison Rae': 86.5,
 'Arctic Monkeys': 80.0,
 'Benson Boone': 89.1,
 'BigXthaPlug': 80.0,
 'Billie Eilish': 92.73,
 'Bruno Mars': 100.0,
 'Chappell Roan': 89.54,
 'Charli xcx': 89.0,
 'Dasha': 89.75,
 'ERNEST': 84.4,
 'Frank Ocean': 87.0,
 'Fuerza Regida': 87.0,
 'Gigi Perez': 77.4,
 'Gracie Abrams': 89.0,
 'Gunna': 85.0,
 'Hanumankind': 91.0,
 'Hozier': 76.7,
 'Jimin': 93.0,
 'Kacey Musgraves': 87.2,
 'Kali Uchis': 89.0,
 'Kalmi': 91.0,
 'Kendrick Lamar': 91.4,
 'Kyle Richh': 62.67,
 'Lady Gaga': 100.0,
 'Luke Combs': 85.0,
 'Mitski': 88.0,
 'Morgan Wallen': 83.77,
 'NLE Choppa': 62.67,
 'Neton Vega': 86.4,
 'Noah Kahan': 88.0,
 'Peso Pluma': 86.4,
 'Playboi Carti': 89.27,
 'Post Malone': 85.2,
 'Sabrina Carpenter': 89.02,
 'Shaboozey': 86.1,
 'Tate McRae': 86.25,
 'Teddy Swims': 90.1,
 'The Creator': 89.0,
 'The Kid LAROI': 88.3,
 'The Marías': 77.0,
 'The Neighbourhood': 91.0,
 'The Weeknd': 88.1,
 'Tommy Richman': 91.0,
 'Travis Scott': 89.2,
 'Tucker Wetmore': 41.5,
 'Tyler': 89.0,
 'Zach Bryan': 85.42,
 'girl in red': 85.5}
```

### Question 7 [20 points]

For this question you will have to go back to the spotify csv file. 

List all the songs that appear in the top list by the season of their release date.
(For this question, you can assume that full months belong to a season, so Winter is Jan,Feb,Mar, Spring is Apr,May,Jun, Summer is Jul,Aug,Sep, and Fall is Oct,Nov,Dec)

Write a **TopSongsPerSeason** function that takes a **filename** and computes a dictionary that for each season gives a list of  (song_name, release_year_month (in yyyy-mm format), average popularity) triplets (as tuples), **sorted by the popularity (decreasing)**. 

The Dictionary should look like this:
{'Summer':[[song_name,release_month,popularity],[song_name,release_month,popularity]]}

```python
#FUNCTION FOOTPRINT
#def TopSongsPerSeason(filename):
# returns dict
from typing import Tuple

def TopSongsPerSeason(filename: str) -> Dict[str, List[Tuple[str, str, float]]]:
    thedict = {}
    thedict['Fall'] = []
    thedict['Spring'] = []
    thedict['Summer'] = []
    thedict['Winter'] = []
    with open(filename) as file:
        newdict = {}
        reader = csv.reader(file)
        next(reader)
        for num, row in enumerate(reader):
            if row[1] not in newdict.keys():
                newdict[row[1]] = [row[8]]
            else:
                newdict[row[1]].append(row[8])
        avgdict = {}
        for i in newdict:
            sum = 0
            for j in newdict[i]:
                sum += float(j)
            avg = sum/len(newdict[i])
            avgdict[i] = avg
    file.close()
    with open(filename) as file:
        reader = csv.reader(file)
        next(reader)
        for num, row in enumerate(reader):
            newstr = row[12].split('-')[0] + '-' + row[12].split('-')[1]
            newtuple = (row[1], newstr, avgdict[row[1]])
            if newstr[-2:] == '01' or newstr[-2:] == '02'or newstr[-2:] == '03':
                if newtuple not in thedict['Winter']:
                    thedict['Winter'].append(newtuple)
            if newstr[-2:] == '04' or newstr[-2:] == '05'or newstr[-2:] == '06':
                if newtuple not in thedict['Spring']:
                    thedict['Spring'].append(newtuple)
            if newstr[-2:] == '07' or newstr[-2:] == '08'or newstr[-2:] == '09':
                if newtuple not in thedict['Summer']:
                    thedict['Summer'].append(newtuple)
            if newstr[-2:] == '10' or newstr[-2:] == '11'or newstr[-2:] == '12':
                if newtuple not in thedict['Fall']:
                    thedict['Fall'].append(newtuple)
    file.close()
    thedict['Winter'] = sorted(thedict['Winter'],key=lambda forfield:(forfield[2], forfield[1], forfield[0]), reverse=True)
    thedict['Fall'] = sorted(thedict['Fall'],key=lambda forfield:(forfield[2], forfield[1], forfield[0]), reverse=True)
    thedict['Spring'] = sorted(thedict['Spring'],key=lambda forfield:(forfield[2], forfield[1], forfield[0]), reverse=True)
    thedict['Summer'] = sorted(thedict['Summer'],key=lambda forfield:(forfield[2], forfield[1], forfield[0]), reverse=True)
    return thedict
    

seasons_songs = TopSongsPerSeason('top_spotify_songs_USA_earlyOct24.csv')
print(seasons_songs)
```

Output:
```
{'Fall': [('NIGHTS LIKE THIS', '2023-11', 88.3), ('Stick Season', '2022-10', 88.0), ('we fell in love in october', '2018-11', 85.5), ('Mmhmm', '2023-12', 80.0), ('Wind Up Missin’ You', '2024-10', 41.5)], 'Spring': [('BIRDS OF A FEATHER', '2024-05', 98.2), ('Good Luck, Babe!', '2024-04', 96.0), ('Not Like Us', '2024-05', 91.4), ('WILDFLOWER', '2024-05', 91.0), ('MILLION DOLLAR BABY', '2024-04', 91.0), ('Sweater Weather', '2013-04', 91.0), ('Beautiful Things', '2024-04', 90.0), ("I Love You, I'm Sorry", '2024-06', 89.0), ('Slow It Down', '2024-04', 88.2), ('Pink Skies', '2024-05', 87.0), ('LA PATRULLA', '2024-06', 86.4), ('A Bar Song (Tipsy)', '2024-05', 86.1), ('Something in the Orange', '2022-04', 86.1), ('one of wun', '2024-05', 85.0), ("Ain't No Love In Oklahoma (From Twisters: The Album)", '2024-05', 85.0), ('Sun to Me', '2022-05', 83.44444444444444), ('505', '2007-04', 80.0), ('No One Noticed', '2024-05', 77.0)], 'Summer': [('Die With A Smile', '2024-08', 100.0), ('Taste', '2024-08
```

```python
from A1sub import TopSongsPerSeason
seasons_songs = TopSongsPerSeason('top_spotify_songs_USA_earlyOct24.csv')
pprint.pprint(seasons_songs)
```

```
{'Fall': [('NIGHTS LIKE THIS', '2023-11', 88.3),
          ('Stick Season', '2022-10', 88.0),
          ('we fell in love in october', '2018-11', 85.5),
          ('Mmhmm', '2023-12', 80.0),
          ('Wind Up Missin’ You', '2024-10', 41.5)],
 'Spring': [('BIRDS OF A FEATHER', '2024-05', 98.2),
            ('Good Luck, Babe!', '2024-04', 96.0),
            ('Not Like Us', '2024-05', 91.4),
            ('WILDFLOWER', '2024-05', 91.0),
            ('Sweater Weather', '2013-04', 91.0),
            ('MILLION DOLLAR BABY', '2024-04', 91.0),
            ('Beautiful Things', '2024-04', 90.0),
            ("I Love You, I'm Sorry", '2024-06', 89.0),
            ('Slow It Down', '2024-04', 88.2),
            ('Pink Skies', '2024-05', 87.0),
            ('LA PATRULLA', '2024-06', 86.4),
            ('A Bar Song (Tipsy)', '2024-05', 86.1),
            ('Something in the Orange', '2022-04', 86.1),
            ('one of wun', '2024-05', 85.0),
            ("Ain't No Love In Oklahoma (From Twisters: The Album)",
             '2024-05',
             85.0),
            ('Sun to Me', '2022-05', 83.44444444444444),
            ('505', '2007-04', 80.0),
            ('No One Noticed', '2024-05', 77.0)],
 'Summer': [('Die With A Smile', '2024-08', 100.0),
            ('Taste', '2024-08', 95.1),
            ('Who', '2024-07', 93.0),
            ('Bed Chem', '2024-08', 92.0),
            ('Big Dawgs', '2024-07', 91.0),
            ('HOT TO GO!', '2023-09', 90.8),
            ('ALL RED', '2024-09', 90.5),
            ('Lose Control', '2023-09', 90.1),
            ('FE!N (feat. Playboi Carti)', '2023-07', 89.2),
            ('See You Again (feat. Kali Uchis)', '2017-07', 89.0),
            ('Good Graces', '2024-08', 89.0),
            ('Guess featuring Billie Eilish', '2024-08', 89.0),
            ('Timeless (with Playboi Carti)', '2024-09', 88.1),
            ('Pink Pony Club', '2023-09', 88.0),
            ('Juno', '2024-08', 88.0),
            ('My Love Mine All Mine', '2023-09', 88.0),
            ('I Remember Everything (feat. Kacey Musgraves)', '2023-08', 87.2),
            ('Pink + White', '2016-08', 87.0),
            ('NEL', '2024-07', 87.0),
            ('Red Wine Supernova', '2023-09', 86.9),
            ('Diet Pepsi', '2024-08', 86.5),
            ("It's ok I'm ok", '2024-09', 86.25),
            ('Casual', '2023-09', 86.0),
            ('I Had Some Help (Feat. Morgan Wallen)', '2024-08', 85.2),
            ('Espresso', '2024-08', 85.0),
            ('Please Please Please', '2024-08', 85.0),
            ('28', '2024-07', 84.7),
            ('Lies Lies Lies', '2024-07', 84.6),
            ('Heading South', '2019-09', 84.1),
            ('Sailor Song', '2024-07', 77.4),
            ('Too Sweet', '2024-08', 76.7),
            ('Or What', '2024-09', 62.666666666666664)],
 'Winter': [("Austin (Boots Stop Workin')", '2024-02', 89.75),
            ('Cowgirls (feat. ERNEST)', '2023-03', 84.4),
            ('Thinkin’ Bout Me', '2023-03', 83.0),
            ('Last Night', '2023-03', 81.625)]}
```

```python
# This code is just for you to double check. 
#And also because that is interesting information from the analysis.
for season in seasons_songs.keys():
    print(str(len(seasons_songs[season])) + " top songs were released in " + season) 
```

```python

```

```python

```